Steroidogenesis and androgen/estrogen signaling pathways are altered in in vitro matured testicular tissues of prepubertal mice

Author Response

The following is the authors’ response to the previous reviews

We would like to thank you again for your thorough review of the manuscript. We have taken all comments into account in the revised version of the manuscript. Please find below our detailed responses to your comments.

Reviewing Editor

The manuscript has been improved, but there are some remaining issues that need to be addressed, as outlined in the reviewers' comments. In particular, please pay attention to Figures 1A and 2A as they appear to be the same. Moreover, the original gel images for Western blots should be made available given the concerns raised by Reviewer #1.

Thank you for your recommendations. We have carefully considered all comments and made the requested revisions to improve the manuscript.

Reviewer #1 (Public Review):

In this manuscript, the authors aimed to compare, from testis tissues at different ages from mice in vivo and after culture, multiple aspects of Leydig cells. These aspects included mRNA levels, proliferation, apoptosis, steroid levels, protein levels, etc. A lot of work was put into this manuscript in terms of experiments, systems, and approaches. The technical aspects of this work may be of interest to labs working on the specific topics of in vitro spermatogenesis for fertility preservation.

Second review:

The authors should be commended for substantial improvement in their manuscript for resubmission.

Thank you very much for this second review and your help to improve this manuscript.

Recommendations For The Authors:

Going forward, the authors would be well-served to put a similar amount of effort on first drafts as well, which would both increase reviewer enthusiasm and reduce reviewer workload to document all the deficiencies! Abstract is much improved, and clearly articulates the point of the study.

We are very grateful for all your constructive comments, which have greatly contributed to the improvement of our manuscript.

1. 54 - replace "could be" with was

“could be” was replaced by “was”

1. 75 - delete "being"

“being” was deleted.

1. 103 - would say "indirectly promotes" since Rhox5 is a transcription factor that presumably activates genes in Sertoli cells whose products then affect neighboring germ cells, either by direct action or by influencing Sertoli cell behavior changes

“indirectly” was added in the sentence.

1. 139, 155, elsewhere - haven't seen dpp italicized before, certainly not the norm

In dpp (days post-partum), “pp” is italicized as it is a Latin word.

1. 265 - delete "found"

“found” was deleted.

1. 263-273 - Is the CYP19 protein referred to encoded by the Cyp19a1 gene (line 263)? Should standardize nomenclature...

The CYP19 protein (aromatase) is indeed encoded by the Cyp19a1 gene. The nomenclature was standardized: “CYP19” was replaced by “CYP19A1” in the entire manuscript.

1. 280 - "homolog" doesn't seem like the right word, as it has a very specific meaning with regards to the evolutionary genetic relatedness of genes. Maybe analog?

“homolog” was replaced by “analog”.

1. 306 - would reword to something like "proportions of seminiferous tubules containing round and elongating spermatids" - the because the tubules don't reach spermatid stages

This sentence was reworded as suggested.

1. 310 - delete "resulted in", unnecessary

“resulted in” was deleted.

1. Why are the images shown in Figures 1A and 2A the same? That seems odd - was that intentional? Curious overall why the data is presented in such a way that it's done twice...

We mistakenly presented immunofluorescence images twice. Duplicate images have been removed. In the modified version of this manuscript, Figure 1A shows 3-HSD immunofluorescence staining in cultures of fresh testicular tissues and in their in vivo counterparts while Figure 1 – figure supplement 1A (not Figure 2A) shows 3-HSD immunofluorescence staining in cultures of frozen/thawed testicular tissues.

1. In all the western blots, the cropping is done awfully close to the bands - why is this? Can full gels be shown in a Supplement? And especially in the westerns in Fig. 5C, esp for CYP17A1, the cropping is unacceptable. This reviewer is wondering whether this is an oversight, or whether there is another band below that one that is being masked? Again, should show whole blot for transparency and to ensure Rigor and Reproducibility.

Full gels are shown in the Supplementary File 2. For CYP17A1, we have shown that only one band of the expected molecular weight is obtained with the antibody (Please see photo below). After this verification, the nitrocellulose membranes were cut at the 55 kDa molecular weight band in order to reveal CYP17A1 expression in the upper part of the membranes and the protein used for normalization in the lower part of the membranes.

1. For all figures, wondering why the font sizes are so disparate? This will need to be addressed before publication so it looks more professional.

All figures have been reworked as requested.

Reviewer #3 (Public Review):

Moutard, Laura, et al. investigated the gene expression and functional aspects of Leydig cells in a cryopreservation/long-term culture system. The authors found that critical genetic markers for Leydig cells were diminished when compared to the in-vivo testis. The testis also showed less androgen production and androgen responsiveness. Although they did not produce normal testosterone concentrations in basal media conditions, the cultured testis still remained highly responsive to gonadotrophin exposure, exhibiting a large increase in androgen production. Even after the hCG-dependent increase in testosterone, genetic markers of Leydig cells remained low, which means there is still a missing factor in the culture media that facilitates proper Leydig cell differentiation. Optimizing this testis culture protocol to help maintain proper Leydig cell differentiation could be useful for future human testis biopsy cultures, which will help preserve fertility and child cancer patients.

Overall, the authors addressed most comments and questions from the previous review. The additional data regarding the necrotic area is helpful for interpreting the quality of the cultures. The authors did not conduct a multiple comparison tests although there are multiple comparisons conducted on for a single dependent variable (Fig 2J, Fig 3F, among many others), however, the addition of this multiple comparison is unlikely to change the conclusions of the paper or the figure and, thus is a minor technical detail in this case.

Thank you very much for this second review and your help to improve this manuscript.

Changes from the earlier version to the revised version

Abstract

Abstract

Children undergoing cancer treatments are at risk for impaired fertility. Cryopreserved prepubertal testicular biopsies could theoretically be later matured in vitro to produce spermatozoa for assisted reproductive technology. A complete in vitro spermatogenesis has been obtained from mouse prepubertal testicular tissue, although with low efficiency. Steroid hormones being essential for the progression of spermatogenesis, the aim of this study was to investigate steroidogenesis and steroid signaling in organotypic cultures. Histological, RT-qPCR, western blot analyses and steroid hormone measurements were performed on in vitro cultured mouse prepubertal testicular tissues and age-matched in vivo controls. Despite a conserved density of Leydig cells after 30 days of culture (D30), transcript levels of adult Leydig cell and steroidogenic markers were decreased. Increased amounts of progesterone and estradiol and reduced androstenedione levels were observed at D30, together with decreased transcript levels of steroid metabolizing genes and steroid target genes. hCG was insufficient to facilitate Leydig cell differentiation, restore steroidogenesis and improve sperm yield. In conclusion, this study reports the failure of adult Leydig cell development and altered steroid production and signaling in tissue cultures. The organotypic culture system will need to be further improved before it can be translated in clinics for childhood cancer survivors.

Introduction

Pediatric cancer treatments such as chemotherapy have recognized toxicity on germline stem cells, which could lead to infertility at adulthood (1,2). Freezing of prepubertal testicular tissue containing spermatogonia is a fertility preservation option proposed to prepubertal boys with cancer prior to highly gonadotoxic treatments. Several fertility restoration approaches, whose aim is to mature cryopreserved tissues in vivo or in vitro in order to produce spermatozoa, are being developed. So far, these approaches have not been clinically validated. In vitro maturation strategies are currently being optimized in animal models. In vitro spermatogenesis could indeed be proposed to patients with testicular localization of residual tumor cells, for whom testicular tissue autografting is not indicated (about 30% of patients with acute leukemia). The organotypic culture procedure, which preserves testicular tissue architecture, microenvironment and cell interactions, has been used successfully to obtain spermatozoa from fresh or frozen/thawed mouse prepubertal testicular tissues (3–6). In addition, viable and fertile mouse offspring have been obtained from in vitro produced spermatozoa by oocyte microinjection (3,4).

Earlier version: It has been previously shown that supplementing organotypic culture media with retinol improves the in vitro production of spermatids and spermatozoa in prepubertal mouse testicular tissues cultured at a gas-liquid interphase (5,7). However, in vitro sperm production still remains a rare event. Transcript levels of the androgen receptor (AR)-regulated gene Rhox5 were decreased at the end of the culture period, suggesting that testosterone production by Leydig cells and/or AR transcriptional activity could be impaired in organotypic cultures (8). A decline in intratesticular testosterone levels has been highlighted in vitro in 35- to 49-day cultures of 5 dpp mouse testes, i.e. after the end of the first wave of spermatogenesis (9). Moreover, our recent transcriptomic analysis reveals a decline in mRNA levels of Cyp17a1, encoding a steroidogenic enzyme, in 4- to 30-day cultures of 6 dpp mouse testes (10). Based on these results, Leydig cell maturation and functionality need to be thoroughly explored in order to identify the molecular mechanism that are deregulated in vitro.

Revised version: It has been previously shown that supplementing organotypic culture media with retinol improves the in vitro production of spermatids and spermatozoa in prepubertal mouse testicular tissues cultured at a gas-liquid interphase (5,7). However, in vitro sperm production still remains a rare event. Transcript levels of the androgen receptor (AR)-regulated gene Rhox5 were decreased at the end of the culture period, suggesting that testosterone production by Leydig cells and/or AR transcriptional activity was impaired in organotypic cultures (8). A decline in intratesticular testosterone levels has been highlighted in vitro in 35- to 49-day cultures of 5 dpp mouse testes, i.e. after the end of the first wave of spermatogenesis (9). Moreover, our recent transcriptomic analysis reveals a decline in mRNA levels of Cyp17a1, encoding a steroidogenic enzyme, in 4- to 30-day cultures of 6 dpp mouse testes (10). Based on these results, Leydig cell maturation and functionality need to be thoroughly explored in order to identify the molecular mechanism that are deregulated in vitro.

Earlier version: Two distinct populations of Leydig cells appear during mouse testicular development: (i) fetal Leydig cells arising during embryonic development and regressing over the first two weeks of postnatal life and (ii) adult Leydig cells appearing around one week after birth. Although they both express StAR (steroidogenic acute regulatory protein), CYP11A1 (P450 side chain cleavage), 3β-HSD1 (3β-hydroxysteroid dehydrogenase type 1) and CYP17A1 (P450 17α-hydroxylase/C17-20 lyase), which are steroidogenic proteins required for androgen production, 17β-HSD3 (17β-hydroxysteroid dehydrogenase type 3), which converts androstenedione to testosterone, is expressed only by adult Leydig cells (11,12). Consequently, the major androgens produced in fetal Leydig cells and adult Leydig cells are androstenedione and testosterone, respectively (13). In rodents, adult Leydig cells develop from undifferentiated stem Leydig cells through two intermediate cells, progenitor Leydig cells and immature Leydig cells (14). All these cells, except stem Leydig cells, express CYP11A1, 3β-HSD1 and CYP17A1 (15,16). 17β-HSD3 begins to be expressed in immature Leydig cells (17). The expression of the androgen metabolizing enzyme SRD5A1 (5α-reductase 1), which is high in progenitor and immature Leydig cells, is being silenced in adult Leydig cells (17). Adult Leydig cells express INSL3 (Insulin-like peptide 3) (18,19) as well as SULT1E1 (estrogen sulfotransferase) (12), which protects Leydig cells and seminiferous tubules against estrogen overstimulation by catalyzing the sulfoconjugation and inactivation of estrogens. Estrogens and androgens can also be inactivated in the testis by 17β-HSD2 (17β-hydroxysteroid dehydrogenase type 2), which converts estradiol to estrone and testosterone to androstenedione (20). Different factors, among which DHH (Desert hedgehog), IGF1 (Insulin-like Growth Factor 1) and LH (Luteinizing hormone) regulate the proliferation and differentiation of the Leydig cell lineage: DHH is required for both proliferation and differentiation of stem Leydig cells (21), IGF1 stimulates the proliferation of progenitor Leydig cells and immature Leydig cells (22) and promotes the maturation of immature Leydig cells into adult Leydig cells (23), and LH is critical to both adult Leydig cell differentiation and their precursor proliferation (24).

Revised version: Two distinct populations of Leydig cells appear during mouse testicular development: (i) fetal Leydig cells arising during embryonic development and regressing over the first two weeks of postnatal life and (ii) adult Leydig cells appearing around one week after birth. Although they both express StAR (steroidogenic acute regulatory protein), CYP11A1 (P450 side chain cleavage), 3β-HSD1 (3β-hydroxysteroid dehydrogenase type 1) and CYP17A1 (P450 17α-hydroxylase/C17-20 lyase), which are steroidogenic proteins required for androgen production, 17β-HSD3 (17β-hydroxysteroid dehydrogenase type 3), which converts androstenedione to testosterone, is expressed only by adult Leydig cells (11,12). Consequently, the major androgens produced in fetal Leydig cells and adult Leydig cells are androstenedione and testosterone, respectively (13). In rodents, adult Leydig cells develop from undifferentiated stem Leydig cells through two intermediate cells, progenitor Leydig cells and immature Leydig cells (14). All these cells, except stem Leydig cells, express CYP11A1, 3β-HSD1 and CYP17A1 (15,16). 17β-HSD3 begins to be expressed in immature Leydig cells (17). The expression of the androgen metabolizing enzyme SRD5A1 (5α-reductase 1), which is high in progenitor and immature Leydig cells, is silenced in adult Leydig cells (17). Adult Leydig cells express INSL3 (Insulin-like peptide 3) (18,19) as well as SULT1E1 (estrogen sulfotransferase) (12), which protects Leydig cells and seminiferous tubules against estrogen overstimulation by catalyzing the sulfoconjugation and inactivation of estrogens. Estrogens and androgens can also be inactivated in the testis by 17β-HSD2 (17β-hydroxysteroid dehydrogenase type 2), which converts estradiol to estrone and testosterone to androstenedione (20). Different factors, among which DHH (Desert hedgehog), IGF1 (Insulin-like Growth Factor 1) and LH (Luteinizing hormone) regulate the proliferation and differentiation of the Leydig cell lineage: DHH is required for both proliferation and differentiation of stem Leydig cells (21), IGF1 stimulates the proliferation of progenitor Leydig cells and immature Leydig cells (22) and promotes the maturation of immature Leydig cells into adult Leydig cells (23), and LH is critical to both adult Leydig cell differentiation and their precursor proliferation (24).

The steroid hormones produced by Leydig cells are essential for the progression of spermatogenesis. Testosterone, synthesized under the control of LH, is essential for many aspects of spermatogenesis, including meiotic progression, spermiogenesis and germ cell survival (25–28). Androgen-binding protein (ABP), produced by Sertoli cells under the regulation of FSH, increases the accumulation of androgens in the seminiferous epithelium and makes them available for binding to intracellular AR (29). Estrogens, derived from the aromatization of androgens by CYP19A1 (aromatase) in somatic and germ cells, have been shown to be important for the long-term maintenance of spermatogenesis in the ArKO mouse (lacking aromatase) and for the progression of normal germ cell development in the ENERKI mouse (estrogen nonresponsive ERα knock-in) (30,31). The expression of aromatase is age-dependent, being mostly in Sertoli cells in immature rat testes and then in Leydig cells during adulthood (32,33). In mouse, aromatase is also present within seminiferous tubules, mainly in spermatids (34).

Earlier version: Androgens act directly on somatic cells (Sertoli cells and peritubular myoid cells) through AR to support germ cell development by regulating the expression of target genes. Rhox5, encoding a transcription factor that promotes germ cell survival (35), and Eppin, encoding a serine protease inhibitor, are target genes of testosterone in Sertoli cells, with androgen response elements (ARE) in their promoters (36). Several estrogen receptors mediate the effects of estrogens in the testis. In rodents, ERα was found in Leydig cells and peritubular myoid cells, ERβ was found in Sertoli cells and some germ cells (spermatogonia, spermatocytes) and GPER was detected in Leydig cells, Sertoli cells and germ cells (spermatocytes, spermatids) (37–41). Faah, encoding a hydrolase that promotes Sertoli cell survival by degrading anandamide, is a direct target gene of estrogens in mature Sertoli cells (42,43). Septin12, containing 1 AR and 2 ERα binding sites in its promoter, has been identified as a potential novel target of the androgen and estrogen receptors in human testicular cells (44).

Revised version: Androgens act directly on somatic cells (Sertoli cells and peritubular myoid cells) through AR to support germ cell development by regulating the expression of target genes. Rhox5, encoding a transcription factor that indirectly promotes germ cell survival (35), and Eppin, encoding a serine protease inhibitor, are target genes of testosterone in Sertoli cells, with androgen response elements (ARE) in their promoters (36). Several estrogen receptors mediate the effects of estrogens in the testis. In rodents, ERα was found in Leydig cells and peritubular myoid cells, ERβ was found in Sertoli cells and some germ cells (spermatogonia, spermatocytes) and GPER was detected in Leydig cells, Sertoli cells and germ cells (spermatocytes, spermatids) (37–41). Faah, encoding a hydrolase that promotes Sertoli cell survival by degrading anandamide, is a direct target gene of estrogens in mature Sertoli cells (42,43). Septin12, containing 1 AR and 2 ERα binding sites in its promoter, has been identified as a potential novel target of the androgen and estrogen receptors in human testicular cells (44).

Since steroid hormones play an essential role in the progression of spermatogenesis, it appears necessary to ensure that their syntheses and mechanisms of action are not altered in in vitro cultured testicular tissues. The aim of the present work was therefore to study Leydig cell maturation, steroidogenesis and androgen/estrogen signaling in a comprehensive manner during the in vitro maturation of mouse prepubertal testicular tissues. The expression of Leydig cell markers, Leydig cell differentiation factors, steroidogenic enzymes, steroid receptors, steroid target genes and steroid contents were quantified in cultured tissues and in their in vivo counterparts. Although our organotypic culture conditions support a complete spermatogenesis, a failure of adult Leydig cell development with impaired steroidogenesis and altered steroid hormone signaling was demonstrated in this study.

Results

Leydig cells are partially mature after 30 days of organotypic culture

We first wondered whether Leydig cell number, survival, proliferation and differentiation could be altered under in vitro conditions. 3β-HSD immunofluorescence staining was performed to detect and quantify Leydig cells during mouse postnatal development and in in vitro cultured fresh testicular tissues (FT). The percentage of Leydig cells in apoptosis or in proliferation was measured after cleaved caspase 3 or Ki67 immunofluorescence staining respectively, and the percentage of Leydig cells expressing AR, which is required for their proliferation and maturation (51), was also determined. The transcript levels of genes necessary for Leydig cell differentiation (Igf1, Dhh) and markers of fetal Leydig cells (Mc2r), progenitor/immature Leydig cells (Srd5a1) or adult Leydig cells (Sult1e1, Insl3) were then assessed by RT-qPCR. Finally, the concentration of INSL3, a marker of Leydig cell maturity, was measured by RIA in organotypic culture media and in testicular homogenates.

After 16 days of culture, the number of Leydig cells per cm² of testicular tissue was not significantly different from 22 dpp, the age-matched in vivo control (Figure 1A-B). However, at this time point, the percentage of apoptotic Leydig cells was increased compared to the in vivo control while the percentage of proliferating Leydig cells was unchanged (Figure 1C-D). In addition, the percentage of Leydig cells expressing AR was similar at D16 and 22 dpp (Figure 1E). Whereas no difference was observed in the mRNA levels of Dhh, Srd5a1 and Sult1e1 between D16 and 22 dpp, Igf1 and Insl3 transcripts levels were respectively higher and lower in 16-day organotypic cultures than in vivo (Figure 1F-J). Mc2r transcripts were barely detected at D16 and 22 dpp as well as at later time points, thereby reflecting the low amount of fetal Leydig cells in prepubertal testicular tissues (data not shown).

Leydig cells are partially mature after 30 days of organotypic culture.

Earlier version: (A) Representative images of 3β-HSD expression by Leydig cells during mouse postnatal development (6 dpp, 22 dpp and 36 dpp) and in in vitro cultured tissues after 16 days of culture (D16) or 30 days (D30). A representative image of a negative control, carried out by omitting the primary antibody, is also shown. Testicular tissue sections were counterstained with Hoechst (blue). Dotted lines delineate seminiferous tubules. Scale: 15 µm. (B) Number of 3β-HSD+ Leydig cells per cm² of testicular tissue during mouse postnatal development (6 dpp, 22 dpp and 36 dpp) and in in vitro cultured tissues (D16 and D30). (C-D) Percentage of Leydig cells (C) in apoptosis (3β-HSD and cleaved caspase 3 positive) or (D) in proliferation (3β-HSD and Ki67 positive) in in vivo and in vitro matured tissues. (E) Percentage of 3β-HSD positive Leydig cells expressing AR in in vivo and in vitro matured tissues. (F-J) Relative mRNA levels of Leydig cell differentiation factors (Igf1, Dhh), progenitor/immature Leydig cell (Srd5a1) and adult Leydig cell markers (Sult1e1, Insl3) (normalized to Gapdh and Actb or to Hsd3b1). (K) Intratesticular concentration of INSL3 (pg/mg of tissue) or in the culture medium (pg/mL).

Revised version: (A) Representative images of 3β-HSD expression by Leydig cells during mouse postnatal development (6 dpp, 22 dpp and 36 dpp) and in in vitro cultured tissues after 16 days of culture (D16) or 30 days (D30). A representative image of a negative control, carried out by omitting the primary antibody, is also shown. Testicular tissue sections were counterstained with Hoechst (blue). Dotted lines delineate seminiferous tubules. Scale: 15 µm. (B) Number of 3β-HSD+ Leydig cells per cm² of testicular tissue during mouse postnatal development (6 dpp, 22 dpp and 36 dpp) and in in vitro cultured tissues (D16 and D30). (C-D) Percentage of Leydig cells (C) in apoptosis (3β-HSD and cleaved caspase 3 positive) or (D) in proliferation (3β-HSD and Ki67 positive) in in vivo and in vitro matured tissues. (E) Percentage of 3β-HSD positive Leydig cells expressing AR in in vivo and in vitro matured tissues. (F-J) Relative mRNA levels of Leydig cell differentiation factors (Igf1, Dhh), progenitor/immature Leydig cell (Srd5a1) and adult Leydig cell markers (Sult1e1, Insl3) (normalized to Gapdh and Actb or to Hsd3b1). (K-L) Intratesticular concentration of INSL3 (pg/mg of tissue, K) or in the culture medium (pg/mL, L).

Data are presented as means ± SEM with n = 4 biological replicates for each group. A value of *P < 0.05 was considered statistically significant. n.d.: not determined (under the detection limit)

FT: Fresh Tissue; CSF: Controlled Slow Freezing.

After 30 days of culture, no significant difference in the number of Leydig cells per cm² of testicular tissue and in the percentages of apoptotic and proliferating Leydig cells was observed compared to 36 dpp, the corresponding in vivo time point (Figure 1A-D). Furthermore, the percentage of Leydig cells expressing AR was comparable after in vitro or in vivo maturation (Figure 1E). The transcript levels of Dhh were also similar in in vitro and in vivo matured tissues at the end of the first spermatogenic wave (Figure 1G). However, a reduction in the mRNA levels of Igf1, Srd5a1 and the two adult Leydig cell markers (Sult1e1, Insl3) was found at D30 in comparison to 36 dpp in vivo controls (Figure 1F, H-J). Moreover, intratesticular INSL3 was below the detection limit (<10 pg/mL) in in vitro matured tissues (Figure 1K) and a significant elevation in the concentration of INSL3 was observed in culture medium from D22 to D30 for FT tissues (Figure 1L).

In summary, while the number and proliferation of Leydig cells are not affected by organotypic culture conditions, their differentiation into mature adult Leydig cells is impaired in vitro.

Controlled slow freezing has no impact on Leydig cell density or state of differentiation before or after organotypic culture

In the clinics, testicular biopsies from prepubertal boys are frozen and stored in liquid nitrogen for later use. In order to assess the impact of freezing/thawing procedures (CSF) on Leydig cell number, survival, proliferation and differentiation in organotypic cultures, we conducted the same analyses as above on 6 dpp CSF mouse testicular tissues and on in vitro matured 6 dpp CSF tissues (Figure 1 -figure supplement 1).

CSF had no impact on the number of Leydig cells per cm² of tissue at 6 dpp as well as after 16 or 30 days of culture (Figure 1 -figure supplement 1A-B). Although the percentage of apoptotic Leydig cells was lower in CSF than in FT tissues at 6 dpp, it was similar in CSF and FT tissues at D16 and D30 (Figure 1 -figure supplement 1C). Furthermore, no change in the percentage of Leydig cells in proliferation or expressing AR was found after CSF (Figure 1 - figure supplement 1D-E). Moreover, Dhh and Srd5a1 mRNA levels were similar in CSF and FT tissues (Figure 1 -figure supplement 1G-H). Igf1 transcript levels were decreased in CSF tissues at D16 (Figure 1 -figure supplement 1F). Furthermore, Sult1e1 mRNA levels were higher in CSF tissues at D30 (Figure 1 -figure supplement 1I) and Insl3 transcript levels were higher in CSF than in FT tissues at D16 and D30 (Figure 1 -figure supplement 1J).

Since the number and the state of differentiation of Leydig cells are rather similar in FT and CSF tissues, it can be concluded that the freezing/thawing procedures are not harmful to these cells.

The expression of several actors of steroidogenesis is affected in organotypic cultures

As the differentiation of Leydig cells is not fully completed in organotypic cultures, we next wanted to know if actors of the steroidogenic pathway show deregulated expression in vitro in comparison to physiological conditions, and thus which steps of the steroid hormone biosynthesis pathway may be impaired. The transcript levels of several genes involved in steroidogenesis were measured by RT-qPCR to highlight a potential deregulation of their expression in cultured tissues (Figure 2A-D, F, H-I). The protein levels of two steroidogenic enzymes, 3β-HSD and CYP17A1, were also quantified by western blot (Figure 2E, G).

The expression of several actors of steroidogenesis is downregulated in 30-day organotypic cultures.

(A-I) Relative mRNA levels of Lhcgr, Star, Cyp11a1, Hsd3b1, Cyp17a1, Hsd17b3 and Hsd17b2 (normalized to Gapdh and Actb or to Hsd3b1) and relative protein levels of 3β-HSD (normalized to ACTB) and CYP17A1 (normalized to 3β-HSD) during mouse postnatal development (6 dpp, 22 dpp and 36 dpp) and in in vitro cultured FT or CSF tissues (D16 and D30). (J) Representative images of CYP17A1 expression during mouse postnatal development (22 dpp and 36 dpp) and in in vitro cultured tissues (D16 and D30). A representative image of a negative control, carried out by omitting the primary antibody, is also shown. Testicular tissue sections were counterstained with Hoechst (blue). Dotted lines delineate seminiferous tubules. Scale: 15 µm.

Data are presented as means ± SEM with n = 4 biological replicates for each group. A value of *P < 0.05 was considered statistically significant.

Controlled slow freezing had no impact on the mRNA levels of all the genes examined at 6 dpp, i.e. before culture (Figure 2A-I). At D16, the mRNA levels of Star, Cyp17a1 and Hsd17b3 were decreased in FT and CSF testicular tissues compared to 22 dpp testes (Figure 2B, F, H). Cyp11a1 transcript levels were also lower in FT tissues at D16 than at 22 dpp (Figure 2C). In contrast, Lhcgr, Hsd3b1 and Hsd17b2 transcript levels remained unchanged at this time point (Figure 2A, D, I). 3β-HSD and CYP17A1 protein levels were not different between D16 and 22 dpp (Figure 2E, G).

The mRNA levels of Cyp11a1, Hsd3b1, Cyp17a1 and Hsd17b2 were lower at D30 in both FT and CSF tissues than in the physiological situation (Figure 2C-D, F, I). Despite a significant decrease in Hsd3b1 mRNA levels at D30, the protein levels of 3β-HSD were not significantly different after 30 days of culture and at 36 dpp (Figure 2E). Western blot experiments however showed that the expression of CYP17A1 was also reduced at the protein level in both FT and CSF tissues at D30 compared to 36 dpp (Figure 2G). This steroidogenic enzyme was still detectable by immunofluorescence in Leydig cells in cultured tissues (Figure 2J). Finally, Lhcgr mRNA levels were found significantly reduced at D30 in FT tissues (Figure 2A), while Star and Hsd17b3 mRNA levels were reduced at D30 in CSF tissues (Figure 2B, H).

Thus, the expression of several genes encoding steroidogenic enzymes is decreased in vitro, notably that of Cyp17a1, necessary for the conversion of progesterone to androstenedione.

Increased intratesticular concentrations of progesterone and estradiol combined with decreased intratesticular concentration of androstenedione in organotypic cultures

We then analyzed the steroid hormone content in in vitro cultured testicular tissues. The concentrations of androstenedione, DHEA and testosterone were measured by LC-MS/MS and the concentrations of progesterone and estradiol were assessed by ELISA in both testicular samples and the culture media (Figure 3). DHEA was below the detection limit (<1 ng/mL) in all the samples examined (data not shown). The intratesticular concentrations of progesterone were significantly increased at D16 and D30 in both FT and CSF tissues compared to their respective in vivo controls (Figure 3A). In contrast, the intratesticular concentrations of androstenedione were significantly decreased at D16 in both FT and CSF tissues and at D30 in FT tissues compared to in vivo controls (Figure 3B). The intratesticular concentrations of testosterone were however not different in in vitro and in vivo matured tissues at both time points (Figure 3C). In addition, the intratesticular concentrations of estradiol were significantly higher at D16 and D30 in FT and CSF tissues than in their respective in vivo controls (Figure 3D).

An increased production of progesterone and estradiol and a decreased production of androstenedione are observed after 16 and 30 days of culture of prepubertal mouse testicular tissues.

Intratesticular concentrations of (A) progesterone, (B) androstenedione, (C) testosterone and (D) estradiol during mouse postnatal development (6 dpp, 22 dpp and 36 dpp) and in in vitro cultured fresh (FT) or frozen/thawed (CSF) tissues (D16 and D30). Steroid concentrations were normalized to protein levels. (E-F) Concentrations of (E) progesterone, androstenedione, testosterone and (F) estradiol in the culture medium of FT tissues. (G) Intratesticular concentrations of total, free and esterified cholesterol normalized to tissue mass.

Data are presented as means ± SEM with n = 4 biological replicates for each group. A value of *P < 0.05 was considered statistically significant.

In the organotypic culture medium, a significant increase in the concentrations of progesterone was observed at D26 only (Figure 3E). Furthermore, a significant elevation in the concentrations of androstenedione was observed from D16 to D30 (Figure 3E). Regarding testosterone, following a decrease in the levels of this androgen in the culture media between D2 and D14, an augmentation was then detected between D26 and D30 (Figure 3E). At D30, the concentrations of both androstenedione and testosterone were higher in the culture media of FT tissues than of CSF tissues (Figure 3 -figure supplement 1). Regarding estradiol, no significant change in its concentrations was observed in the media between D6 and D30 (Figure 3F).

Since steroids are derived from cholesterol, intratesticular levels of total, free and esterified cholesterol were also analyzed (Figure 3G). The intratesticular levels of esterified cholesterol were comparable between the different experimental conditions (Figure 3G). However, total cholesterol levels were higher in FT and CSF cultured tissues at D16 and D30 compared to their corresponding in vivo time points (Figure 3G). Moreover, free cholesterol levels were increased in FT and CSF cultured tissues at D30 compared to 36 dpp and in FT tissues at D16 (Figure 3G).

Overall, our data show that the steroid content is altered in cultured testicular tissues, with an excess of free cholesterol, progesterone and estradiol and a deficiency of androstenedione.

Androgen and estrogen signaling are altered in 30-day organotypic cultures

Earlier version: The impact of imbalanced steroid hormone levels on downstream target gene expression was then investigated. The mRNA levels of actors of androgen and estrogen signaling were measured by RT-qPCR (Figures 4 and 5). The protein expression of the androgen receptor (AR), aromatase (CYP19) and fatty acid amide hydrolase (FAAH whose expression is regulated by estrogen) was also assessed by western blot (Figures 4 and 5).

Revised version: The impact of imbalanced steroid hormone levels on downstream target gene expression was then investigated. The mRNA levels of actors of androgen and estrogen signaling were measured by RT-qPCR (Figures 4 and 5). The protein expression of the androgen receptor (AR), aromatase (CYP19A1) and fatty acid amide hydrolase (FAAH whose expression is regulated by estrogen) was also assessed by western blot (Figures 4 and 5).

Androgen signaling is altered in 30-day organotypic cultures.

(A) Relative mRNA levels of Ar (normalized to Gapdh and Actb) and (B) relative protein levels of AR (normalized to ACTB) during mouse postnatal development (6 dpp, 22 dpp and 36 dpp) and in in vitro cultured FT or CSF tissues (D16 and D30). (C) Representative images of AR expression at 22 dpp and 36 dpp and at corresponding in vitro time points (D16 and D30). A representative image of a negative control is shown in Figure 2J (same secondary antibody as for CYP17A1). Testicular tissue sections were counterstained with Hoechst (blue). Solid arrowheads: peritubular myoid cells. Open arrowheads: Sertoli cells. Arrows: Leydig cells. Dotted lines delineate seminiferous tubules (ST). Scale: 15 µm. (D-G) Relative mRNA levels of Rhox5, Septin12, Eppin and Abp (normalized to Gapdh and Actb).

Data are presented as means ± SEM with n = 4 biological replicates for each group. A value of *P < 0.05 was considered statistically significant.

The expression of aromatase and estrogen signaling are impaired after 30 days of organotypic culture.

(A) Relative mRNA levels of Cyp19a1 (normalized to Gapdh and Actb) and (B) relative protein levels of CYP19A1 (normalized to ACTB) during mouse postnatal development (6 dpp, 22 dpp and 36 dpp) and in in vitro cultured FT or CSF tissues (D16 and D30). The bands on the western blot correspond to different isoforms of CYP19. (C) Representative images of CYP19A1 expression at 6 dpp, 22 dpp and 36 dpp and at corresponding in vitro time points (D16 and D30). A representative image of a negative control, carried out by omitting the primary antibody, is also shown. Testicular tissue sections were counterstained with Hematoxylin. Scale: 15 µm. Arrow: Leydig cells. Arrowheads: elongated spermatids. (D) Aromatase activity (normalized to tissue weight). (E-G) Relative mRNA levels of Esr1, Esr2 and Gper1 (normalized to Gapdh and Actb). (H) Relative mRNA levels of Faah (normalized to Gapdh and Actb) and (I) relative protein levels of FAAH (normalized to ACTB). The second band at 80 kDa is an isoform of FAAH (Q8BRM1, UniProtKB).

Data are presented as means ± SEM with n = 4 biological replicates for each group. A value of *P < 0.05 was considered statistically significant.

The transcript levels of Ar, encoding the androgen receptor, were higher at D16 in FT tissues than at 22 dpp but were not significantly different between D30 and 36 dpp (Figure 4A). Conversely, AR protein levels were comparable at D16 in FT tissues and 22 dpp but were more elevated at D30 in both FT and CSF tissues than at 36 dpp (Figure 4B). Our immunofluorescence data showed that AR was expressed in Sertoli cells, peritubular myoid cells and Leydig cells in 16-day and 30-day cultured tissues as well as in the corresponding in vivo controls (Figure 4C). The transcript levels of two AR-regulated genes (Rhox5, Eppin) and the AR/Erα-regulated gene Septin12 were decreased in FT and CSF tissues cultured for 16 days compared to 22 dpp (Figure 4D-F). At D30, Rhox5 and Septin12 mRNA levels were still lower than in vivo while Eppin mRNA levels were higher than at 36 dpp (Figure 4D-F). In addition, a significant increase in Abp mRNA levels was found at D16 and D30 in CSF but not in FT tissues (Figure 4G).

Earlier version: Intriguingly, the expression of Cyp19a1, encoding aromatase, was drastically diminished after 30 days of culture in FT and CSF tissues compared to 36 dpp (Figure 5A). The protein levels of CYP19 were found reduced in both FT and CSF tissues at D16 and D30 (Figure 5B). CYP19 was expressed in Leydig cells at 6 dpp and was then found in elongated spermatids and spermatozoa at 36 dpp, whereas no expression could be detected in cultured testes (Figure 5C). CYP19 enzymatic activity was also significantly decreased in CSF cultures at D30 (Figure 5D). In addition, the transcript levels of the three genes encoding estrogen receptors (Esr1, Esr2, Gper1) and of the estrogen target gene Faah were significantly lower at D30 than at 36 dpp (Figure 5E-H). Esr1 mRNA levels were also decreased at D16 (Figure 5E). FAAH protein levels were decreased in 16-day FT and 30-day CSF organotypic cultures (Figure 5I).

Revised version: Intriguingly, the expression of Cyp19a1, encoding aromatase, was drastically diminished after 30 days of culture in FT and CSF tissues compared to 36 dpp (Figure 5A). The protein levels of CYP19A1 were reduced in both FT and CSF tissues at D16 and D30 (Figure 5B). CYP19A1 was expressed in Leydig cells at 6 dpp and in elongated spermatids and spermatozoa at 36 dpp, whereas no expression could be detected in cultured testes (Figure 5C). CYP19A1 enzymatic activity was also significantly decreased in CSF cultures at D30 (Figure 5D). In addition, the transcript levels of the three genes encoding estrogen receptors (Esr1, Esr2, Gper1) and of the estrogen target gene Faah were significantly lower at D30 than at 36 dpp (Figure 5E-H). Esr1 mRNA levels were also decreased at D16 (Figure 5E). FAAH protein levels were decreased in 16-day FT and 30-day CSF organotypic cultures (Figure 5I). Thus, the expression of many genes in the androgen and estrogen signaling pathways are deregulated under organotypic culture conditions.

Removed in the revised version: Thus, the expression of many genes in the androgen and estrogen signaling pathways are deregulated under organotypic culture conditions.

hCG is insufficient to stimulate Leydig cell differentiation and to restore steroidogenesis and steroid hormone signaling

Earlier version: As LH is known to stimulate the differentiation of the adult Leydig cell lineage and steroidogenesis, we wondered whether supplementation of organotypic culture media with hCG (LH homolog) could restore steroidogenesis to control levels. In order to determine the dose to be used, we examined the impact of different hCG concentrations (5 pM, 50 pM, 1 nM, 5 nM or 50 nM) on the intratesticular concentrations of testosterone and the secretion of this hormone in the organotypic culture medium. Since the concentration of hCG leading to the best response (i.e. plateau of the dose-response curves for intratesticular and secreted testosterone) was 1 nM, we used this dose for the rest of the analyses (Figure 6 -figure supplement 1).

Revised version: As LH is known to stimulate the differentiation of the adult Leydig cell lineage and steroidogenesis, we wondered whether supplementation of organotypic culture media with hCG (LH analog) could restore steroidogenesis to control levels. In order to determine the dose to be used, we examined the impact of different hCG concentrations (5 pM, 50 pM, 1 nM, 5 nM or 50 nM) on the intratesticular concentrations of testosterone and the secretion of this hormone in the organotypic culture medium. Since the concentration of hCG leading to the best response (i.e. plateau of the dose-response curves for intratesticular and secreted testosterone) was 1 nM, we used this dose for the rest of the analyses (Figure 6 -figure supplement 1).

Upon supplementation with 1 nM hCG, intratesticular androstenedione levels were significantly increased at D30 in FT but not CSF tissues (Figure 6A). The intratesticular testosterone levels were higher at D30 in both FT and CSF tissues following hCG supplementation (Figure 6B). Moreover, a significant increase in androstenedione and testosterone concentrations was observed in the culture media of both FT and CSF tissues following hCG supplementation (Figure 6C-D). The amounts of androgens released into the culture media were however lower for CSF tissues than for FT tissues (Figure 6C-D).

hCG is insufficient to stimulate Leydig cell differentiation and restore steroidogenesis and steroid hormone signaling.

(A-D) Levels of androgens (androstenedione and testosterone) (A-B) in testicular tissues and (C-D) in culture medium with 1 nM hCG or without supplementation. (E-S) Relative mRNA levels of genes involved in steroidogenesis and androgen/estrogen signaling pathways (normalized to Gapdh and Actb or to Hsd3b1) with or without hCG supplementation. (T) Proportion of seminiferous tubules containing the most advanced type of germ cells after in vitro culture with or without hCG. P < 0.05: * vs 36 dpp and # vs D30 FT. (U) Sertoli cell number in seminiferous tubules after in vitro culture with or without hCG.

Earlier version: Data are presented as means ± SEM with n = 4 biological replicates for each group. A value of *P < 0.05 was considered statistically significant.

Revised version: Data are presented as means ± SEM with n = 4 biological replicates for each group. A value of *P < 0.05 was considered statistically significant. FT: Fresh Tissue; CSF: Controlled Slow Freezing.

Removed in the revised version: FT: Fresh Tissue; CSF: Controlled Slow Freezing.

The addition of 1 nM hCG had no impact on the transcript levels of Star, Hsd17b3, Hsd3b1, Ar and Esr1 in both FT and CSF cultures (Figure 6F, I, J, L, P). Hsd17b3 and Esr1 mRNA levels still remained lower upon hCG supplementation than in vivo (Figure 6I, P). No effect of hCG was also observed in FT cultures for Cyp17a1, Faah, Srd5a1, Sult1e1, Insl3, Igf1, Dhh, Abp and Hsd17b2 mRNA levels (Figure 6I, J, S, Figure 6 -figure supplement 2). Lhcgr, Cyp11a1, Cyp19a1, Esr2 and Gper1 mRNA levels were increased after hCG supplementation in FT cultures (Figure 6E, G-H, K, Q, R). Whereas Cyp11a1 and Cyp19a1 mRNA levels were still lower after hCG supplementation than in vivo (Figure 6G, K), Esr2 and Gper1 transcripts were restored to their physiological levels (Figure 6Q-R). In contrast, mRNA levels of Rhox5, Eppin, Septin12 and Igf1 were decreased after hCG supplementation in FT cultures, with Rhox5 and Septin12 transcript levels lower than those observed in vivo (Figure 6M-O).

Earlier version: Finally, the mean proportions of seminiferous tubules reaching the round spermatid and elongated spermatid stages were lower in tissues cultured with than without hCG (Figure 6T). The number of Sertoli cells inside seminiferous tubules was unchanged following hCG supplementation, but was higher than in in vivo controls (Figure 6U).

Revised version: Finally, the mean proportions of seminiferous tubules containing round and elongated spermatids were lower in tissues cultured with than without hCG (Figure 6T). The number of Sertoli cells inside seminiferous tubules was unchanged following hCG supplementation, but was higher than in in vivo controls (Figure 6U).

Earlier version: Although the addition of hCG resulted in increased androgen levels, the expression of adult Leydig cell markers, of several steroidogenic genes and of steroid hormone-regulated genes remained low, thus not promoting the progression of in vitro spermatogenesis.

Revised version: Although the addition of hCG increased androgen levels, the expression of adult Leydig cell markers, of several steroidogenic genes and of steroid hormone-regulated genes remained low, thus not promoting the progression of in vitro spermatogenesis.

Discussion

The present study shows for the first time and in a comprehensive manner that Leydig cell maturation and functionality as well as steroid hormone signaling are impaired in cultures of fresh and frozen/thawed immature mouse testes. Only few differences were found between fresh and frozen/thawed in vitro matured tissues.

A similar density of Leydig cells was found in organotypic cultures and in corresponding in vivo controls. However, Leydig cells only partially matured in vitro. At D16, the expression of the Leydig cell differentiation factors Igf1 and Dhh was unaffected. At D30, the proliferation/apoptosis balance was undisturbed in Leydig cells while the expression of Insl3 was faint, as well as that of Sult1e1, another Leydig cell maturity marker. The deficient Leydig cell maturation observed at D30 could be related to the decreased levels of Igf1, which encodes insulin-like growth factor 1 known to promote the maturation of immature Leydig cells into adult Leydig cells (23).

Leydig cell functionality was also altered in organotypic cultures. Indeed, the expression of several genes encoding steroidogenic enzymes (Cyp11a1, Cyp17a1, Hsd17b3, Hsd17b2) was downregulated in these cells (Figure 7). Reduced Cyp17a1 mRNA levels in organotypic cultures were also recently highlighted using a bulk RNA-seq approach (10). Decreased Cyp11a1 and Cyp17a1 transcript levels were previously reported in Igf1-/- mice, together with decreased intratesticular testosterone levels (22). Furthermore, deficiency in insulin-like growth factors signaling in Insr-/-/Igf1r-/- mouse Leydig cells was shown to impair testicular steroidogenesis (decreased Lhcgr, Star, Cyp11a1, Cyp17a1, Hsd17b3, Srd5a1 and Insl3 mRNA levels) and increase estradiol production (52). Here, we report for the first time an accumulation of estradiol in in vitro cultured tissues, which could be deleterious for sperm production. To promote Leydig cell maturation and lower intratesticular estradiol levels in vitro, supplementation of culture medium with IGF1 could thus be later envisaged, and even more so since this factor has been shown to increase, although modestly, the percentages of round and elongated spermatids in cultured mouse testicular fragments (53). To reduce estrogen levels and increase sperm production, supplementation of the organotypic culture medium with selective estrogen receptor modulators such as tamoxifen or with aromatase inhibitors such as letrozole could also be considered in future studies. These molecules have indeed proven to be effective in increasing sperm concentration in infertile men (54–56).

Altered steroidogenesis and androgen/estrogen signaling pathways in in vitro matured testicular tissues of prepubertal mice

A similar density of Leydig cells is found after 30 days of organotypic culture (D30) and at 36 days postpartum, the corresponding in vivo time point. However, Leydig cells are partially mature in vitro with a decrease in Sult1e1 and Insl3 mRNA levels (adult Leydig cell markers). The mRNA levels of Cyp11a1, Cyp17a1 and Hsd17b3 encoding steroidogenic enzymes and the protein levels of CYP17A1 are decreased in vitro. Increased amounts of cholesterol, progesterone and estradiol and decreased androstenedione intratesticular levels are observed at D30. Furthermore, despite testosterone levels similar to in vivo, the expression of the androgen receptor (AR) and of the androgen binding protein (Abp), androgen signaling is altered at D30, with decreased transcript levels of the androgen target gene Rhox5 and of Septin12. Moreover, with decreased expression and activity of aromatase and decreased estrogen receptor expression, estrogen signaling is impaired at D30, leading to decreased transcript and protein levels of the estrogen target gene Faah.

Our work also revealed an elevation in progesterone and a reduction in androstenedione in in vitro matured tissues (Figure 7), which could arise from the reduced expression of Cyp17a1. A greater sperm production has been highlighted in PRKO mice (lacking the progesterone receptor), thereby showing the inhibitory action of progesterone on spermatogenesis (57). The cumulative excess of progesterone and estradiol in cultured testicular tissues may therefore slow the progression of in vitro spermatogenesis and lead to a poor sperm yield.

Despite the disturbed steroidogenic activity of Leydig cells in vitro, the accumulation of estradiol and progesterone, the decrease in androstenedione, and in contradiction with a previous study (9), intratesticular testosterone levels in cultured tissues were not significantly different from in vivo. Moreover, a complete in vitro spermatogenesis was achieved, albeit with a low efficiency as previously described (5,6). We also found that the percentage of Leydig cells expressing AR was unchanged under in vitro conditions. The increased AR protein levels at D30 could thus be the result of the higher number of Sertoli cells within seminiferous tubules. Intriguingly, common points can be observed in our organotypic cultures and in LCARKO (Leydig cell-specific Ar-/-) mice (58): decreased Insl3, Cyp17a1 and Hsd17b3 mRNA levels, increased intratesticular progesterone levels and unchanged intratesticular testosterone levels compared to controls. This could suggest that AR signaling is dysfunctional in Leydig cells in vitro. The downregulation of Rhox5 expression (Figure 7), which confirms our previous findings (8), further suggests that AR signaling may also be impaired in Sertoli cells. Recently, in a mouse model with disrupted dimerization of the AR ligand-binding-domain, a decreased expression of Hsd17b3 and of AR-regulated genes, such as Rhox5 and Insl3, a higher percentage of Sertoli cells and a decreased production of round and elongated spermatids were found (59), which further supports our hypothesis that AR signaling may be defective in organotypic cultures. In the SPARKI (SPecificity-affecting AR KnockIn)-AR mouse model, which has a defect in binding to AR-specific DNA motifs while displaying normal interaction with the classical AREs, it has been shown that Rhox5 expression is regulated by an AR-selective ARE, while Eppin expression is regulated by a classical, nonselective ARE (60). The differential regulation of these two androgen-responsive genes could explain why Rhox5 mRNA levels were downregulated while Eppin mRNA levels were upregulated in our cultures. We also report here that the transcript levels of Septin12, a potential target of the androgen and estrogen receptors, was reduced in vitro (Figure 7).

Additionally, estrogen signaling was impaired in organotypic cultures, with a low expression of the estrogen-synthesizing enzyme aromatase, of estrogen receptors and of the estrogen target gene Faah (Figure 7). Since the expression of Cyp19a1, Esr1 and Esr2 (encoding aromatase, ERα and ERβ, respectively) can be downregulated by estradiol in the rat testis (61,62), we hypothesize that the transcription of these genes may be negatively controlled by elevated estradiol levels in our tissue cultures. Faah, whose promoter activity engages ERβ and the histone demethylase LSD1, is a direct target gene of estrogens in Sertoli cells (42). The Faah proximal promoter is not regulated by estrogen in immature Sertoli cells, as they express ERβ at the same level as mature cells but do not express LSD1 (42). Thus, the decreased Faah transcript levels in organotypic cultures could be the consequence of the downregulated ERβ expression and/or of the immaturity of Sertoli cells. Whether Sertoli cells in organotypic cultures are as mature as in vivo is unknown and will be the focus of future studies. The transcription of Insl3, which is repressed by estradiol in a MA-10 Leydig cell line model (63), could also be inhibited by the excess of estradiol in our tissue cultures. Interestingly, our in vitro cultured testes exhibit common features with the testes of Sult1e1-/- mice lacking estrogen sulfotransferase, an enzyme that catalyzes the sulfoconjugation and inactivation of estrogens (64): decreased Cyp17a1 expression as well as increased progesterone levels and local estrogen activity (65,66). In organotypic cultures, estrogen excess could be explained by low Sult1e1 expression. In addition, reduced transcript levels of Hsd17b2, encoding an enzyme that converts estradiol to estrone and testosterone to androstenedione, may also explain why estradiol levels remain elevated in cultures while testosterone levels are similar to controls and androstenedione levels are low.

Earlier version: Finally, as previously reported (5,6), the presence of LH or its homolog hCG was not necessary to reconstitute a full first spermatogenic wave in vitro. However, plasma LH concentration is known to rise in vivo from 21-28 dpp in the mouse model (67). We previously showed that the addition of 0.075 nM (50 IU/L) hCG together with FSH from D7 onwards modestly increased the number of spermatozoa produced in organotypic cultures (5). To mimic as closely as possible the physiological conditions, 1 nM hCG was added in the culture medium from D16 onwards. Even though Leydig cells responded to this stimulation by synthesizing and secreting more androgens, supplementation of the culture medium with 1 nM hCG alone was not sufficient to facilitate Leydig cell differentiation, restore steroidogenesis and steroid hormone signaling and increase the yield of in vitro spermatogenesis.

Revised version: Finally, as previously reported (5,6), the presence of LH or its analog hCG was not necessary to reconstitute a full first spermatogenic wave in vitro. However, plasma LH concentration is known to rise in vivo from 21-28 dpp in the mouse model (67). We previously showed that the addition of 0.075 nM (50 IU/L) hCG together with FSH from D7 onwards modestly increased the number of spermatozoa produced in organotypic cultures (5). To mimic as closely as possible the physiological conditions, 1 nM hCG was added in the culture medium from D16 onwards. Even though Leydig cells responded to this stimulation by synthesizing and secreting more androgens, supplementation of the culture medium with 1 nM hCG alone was not sufficient to facilitate Leydig cell differentiation, restore steroidogenesis and steroid hormone signaling and increase the yield of in vitro spermatogenesis.

In conclusion, the present study shows the partial maturation and the disturbed steroidogenic activity of Leydig cells, the abnormal steroid hormone content as well as the altered androgen and estrogen signaling in organotypic cultures of fresh and frozen/thawed prepubertal mouse testicular tissues. Altogether, these defects could contribute to the low efficiency of in vitro spermatogenesis. It will therefore be necessary to optimize the culture medium by adding factors that promote proper Leydig cell differentiation as well as the culture design to prevent central necrosis from affecting the survival and/or the growth and/or the differentiation of the testis in culture. The development of an optimal model of in vitro spermatogenesis could be useful in advancing the field of fertility restoration, but could also have wider applications, as it could be useful for studying the initiation of puberty as well as the impact of cancer therapies, drugs, chemicals and environmental agents (e.g. endocrine disruptors) on the developing testis.

Materials and methods

Key resources table

Ethical approval

All the experimental procedures were approved by the Institutional Animal Care and Use Committee of Rouen Normandy University under the licence/protocol number APAFiS #38239.

Mice and testis collection

CD-1 mice (Charles River Laboratories, L’Arbresle, France) were housed in a temperature-controlled room (22–23°C) under a 12-h light/dark cycle. Prepubertal 6-day postpartum (dpp) male mice were euthanized by decapitation and underwent a bilateral orchiectomy. Testes were transferred to Petri dishes containing α-MEM without phenol red (Gibco by Life Technologies, Saint-Aubin, France) and the complete removal of the tunica albuginea was performed with two needles under a binocular magnifier (S8AP0, Leica Microsystems GmbH, Wetzlar, Germany). Testes were then either cultured immediately (culture from fresh tissues), or after a freezing/thawing cycle (Figure 8). Moreover, mice aged 22 and 36 dpp were euthanized by CO2 asphyxiation and their testes were used as in vivo controls for 16 and 30 days of culture, respectively (Figure 8).

Scheme of the study design.

At 6 dpp (D0), the testis contains only spermatogonia and the initiation of meiosis occurs between 7 and 9 dpp. At 22 dpp (D16), meiosis ends and the first round spermatids appear. At 36 dpp (D30), the first spermatozoa appear and this is the end of the first spermatogenic wave. Arrows represent times at which the medium was collected (D0, D2, D6, D10, D14, D18, D22, D26 and D30).

BM: Basal Medium; CSF: Controlled Slow Freezing; dpp: days postpartum; hCG: human Chorionic Gonadotropin; KSR: KnockOut Serum replacement; PR-: without phenol red; Rol: retinol

Controlled slow freezing (CSF) and thawing of testicular tissues

Freezing procedure

Testes were placed into cryovials (Dominique Dutscher, Brumath, France) containing 1.5 mL of the following cryoprotective medium: Leibovitz L15 medium (Eurobio, Courtabœuf, France) supplemented with 1.5 M dimethylsulfoxyde (DMSO, Sigma-Aldrich, Saint-Quentin Fallavier, France), 0.05 M sucrose (Sigma-Aldrich), 10% (v/v) fetal calf serum (FCS, Life Technologies) and 3.4 mM vitamin E (Sigma-Aldrich) (45). After a 30-min equilibration at 4°C, samples were frozen in a programmable freezer (Nano Digitcool, CryoBioSystem, L’Aigle, France) with a CSF protocol: start at 5°C, then -2°C/min until reaching -9°C, stabilization at -9°C for 7 min, then -0.3°C/min until -40°C and -10°C/min down to -140°C. Testicular tissues were then plunged and stored in liquid nitrogen.

Thawing procedure

Cryotubes were warmed for 1 min at room temperature (RT) and then for 3 min in a water bath at 30°C. They were then successively incubated at 4°C in solutions containing decreasing concentrations of cryoprotectants for 5 min each [solution 1: 1 M DMSO, 0.05 M sucrose, 10% FCS, 3.4 mM vitamin E, Leibovitz L15; solution 2: 0.5 M DMSO, 0.05 M sucrose, 10% FCS, 3.4 mM vitamin E, Leibovitz L15; solution 3: 0.05 M sucrose, 10% FCS, 3.4 mM vitamin E, Leibovitz L15; solution 4: 10% FCS, 3.4 mM vitamin E, Leibovitz L15].

Organotypic cultures at a gas–liquid interphase

In vitro tissue cultures were performed as previously described (3,5). Briefly, prepubertal 6-day old mouse testes, which contain spermatogonia as the most advanced type of germ cells, were first cut into four fragments (approximately 0.75 mm3 each, which was previously determined to be the most appropriate size for mouse in vitro spermatogenesis) (7). They were placed on top of two 1.5% (w/v) agarose Type I gels (Sigma-Aldrich) half-soaked in medium. Testicular tissues were then cultured under 5% CO2 at 34°C for 16 days (D16), which corresponds to the end of meiosis and the appearance of the first round spermatids, or 30 days (D30) to explore the end of the first spermatogenic wave. No growth of the fragments was observed during the organotypic culture period (Figure 8-figure supplement 1A). The basal medium (BM) contained α-MEM without phenol red, 10% KSR (KnockOut Serum Replacement, Gibco by Life Technologies), 0.1 mg/mL streptomycin and 100 UI/mL penicillin (Sigma-Aldrich). The medium was chosen without phenol red because of its estrogen activity (46). Retinol (10-6 M, Sigma-Aldrich) was added in all organotypic cultures from D2 and then every 8 days in order to respect the meiosis entry cycle of spermatogonia (5). Furthermore, the basal medium was supplemented or not with 1 nM hCG (human Chorionic Gonadotropin, MSD France, Courbevoie, France) from D16 in order to stimulate Leydig cell differentiation and assess Leydig cell functionality. Media were prepared just before use and were replaced twice a week.

Histological analyses

Tissue fixation, processing and sectioning

Testicular tissues were fixed with Bouin’s solution (Sigma-Aldrich) or 4% paraformaldehyde (PFA, Sigma-Aldrich) for 2 h at room temperature. They were then dehydrated in ethanol in the Citadel 2000 tissue processor (Shandon, Cheshire, UK) and embedded in paraffin. Tissue sections (3 μm thick) were prepared with the RM2125 RTS microtome (Leica) and were mounted on Polysine slides (Thermo Fisher Scientific, Waltham, MA, USA).

Periodic Acid Schiff (PAS) reaction

A PAS reaction was then performed on in vitro matured tissues. Tissue sections were deparaffinized in xylene and rehydrated in decreasing concentrations of ethanol. Slides were then immersed in 0.5% periodic acid (Thermo Fisher Scientific) for 10 min, rinsed for 5 min in water and then placed for 30 min in the Schiff reagent (Merck, Darmstadt, Germany). After three 2-min washes in water, sections were counterstained with Mayer’s hematoxylin and mounted with Eukitt (CML, Nemours, France). Images were acquired on a DM4000B microscope (Leica) at a ×400 magnification. The most advanced type of germ cells present in the testicular fragments was analyzed in at least 30 cross-sectioned seminiferous tubules (located at the periphery of the cultured fragments, i.e. outside of the necrotic region) (Figure 8-figure supplement 1E-F) from 2 sections with the Application Suite Core v2.4 software (Leica).

Immunofluorescence staining

After deparaffinization, rehydration and a 3-min wash in phosphate buffered saline (PBS, Sigma-Aldrich), tissue sections were boiled in 10 mM citrate buffer pH 6.0 (Diapath, Martinengo, Italy) for 40 min at 96°C. They were cooled for 20 min at RT and rinsed in distilled water for 5 min. A permeabilization step with 0.1% (v/v) Triton X-100 (Sigma-Aldrich) was performed at RT for 15 min for Ki67/3β-HSD immunostaining. Non-specific sites were blocked with 5% (w/v) bovine serum albumin (BSA, Sigma-Aldrich) and 5% (v/v) horse serum (Sigma-Aldrich). Slides were then incubated in humidified environment with primary antibodies (Supplementary File 1a), rinsed 3 times in PBST (PBS with 0.05% Tween-20) and incubated with appropriate secondary antibodies (Supplementary File 1a). For Ki67/3β-HSD, a sequential protocol was performed as follows: incubation with anti-Ki67 antibodies overnight at 4°C, incubation with secondary antibodies coupled to Alexa 594 for 60 min at RT, fixation with 4% PFA for 15 min at RT and incubation with anti-3β-HSD antibodies directly coupled to Alexa 488 for 90 min at RT. Negative controls were carried out by omitting the primary antibodies. Sections were washed, dehydrated with ethanol and mounted in Vectashield with Hoechst. Images were acquired on a THUNDER Imager 3D Tissue microscope (Leica) at a ×400 magnification. Leydig cell number (3β-HSD+) was normalized to tissue area (cm²).

Immunohistochemical staining

After deparaffinization and rehydration, endogenous peroxidases were blocked with HP Block (Dako, Les Ulis, France) for 30 min and non-specific binding sites were blocked with Ultra-V Block solution (Thermo Fisher Scientific) for 10 min at RT. Tissue sections were then incubated overnight at 4°C with anti-CYP19A1 antibodies (Supplementary File 1a). After three 5-min washes in PBS, they were incubated for 10 min at RT with biotinylated polyvalent secondary antibodies (UltraVision Detection System HRP kit, Thermo Fisher Scientific). After three 5-min washes in PBS, a 10-min incubation at RT with streptavidin associated with peroxidase (UltraVision Detection System HRP kit, Thermo Fisher Scientific) was performed. The labeling was revealed after application of a chromogenic substrate (EnVision FLEX HRP Magenta Chromogen, Dako) for 10 min at RT. A negative control was carried out by omitting the primary antibody.

RNA extraction and RT-qPCR

RNA extraction

Total RNA was extracted from testicular samples using RNeasy Micro kit (Qiagen, Courtabœuf, France) according to the manufacturer’s instructions. For in vitro cultured tissues, the central necrotic area (16-27% of the explants, Figure 8-figure supplement 1B-F) was carefully removed before RNA extraction, so that transcript levels in the healthy part of the samples (i.e. where in vitro spermatogenesis occurs) could be measured and compared with in vivo controls. To avoid contamination with genomic DNA, extracted RNA was incubated with two units of TURBO DNase (Life Technologies) for 45 min at 37°C. The amount of the RNA samples was measured with a NanoDrop Spectrophotometer (NanoDrop Technologies, Wilmington, DE, USA) and purity was determined by calculating the ratio of optical densities at 260 nm and 280 nm.

Reverse transcription

The reverse transcription reaction was performed in a 96-well plate from 1 μg total RNA with 4 μL of qScript cDNA SuperMix 5X (QuantaBio, Gaithersburg, MD, USA) and ribonuclease-free water to adjust the volume to 20 μL. Reverse transcription was performed according to the following program: 5 min at 25°C, 30 min at 42°C and 5 min at 85°C. The complementary DNA (cDNA) obtained were diluted 1:10.

Polymerase chain reaction

cDNA amplifications were carried out in a total volume of 13 μL containing 6 ng of cDNA templates, 6.5 μL of SYBR Green (Thermo Fisher Scientific) and 300 nM of each primer. Specific primers are listed in Supplementary File 1b. Samples were dispensed using the Bravo pipetting robot (Agilent Technologies, Santa Clara, CA, USA). Reactions were performed in 384-well plates (Life Technologies) in a Quant Studio 12K Flex system. The amplification condition was 20 s at 95°C followed by 40 cycles (1 s at 95°C, 20 s at 60°C) and a final step of denaturation of 15 s at 95°C, 1 min at 60°C and 15 s at 95°C. Melting curves were obtained to ensure the specificity of PCR amplifications. The size of the amplicons was verified by agarose gel electrophoresis (E-gel 4%, Life Technologies). The relative expression level of each gene was normalized to two housekeeping genes as recommended by the MIQE guidelines (47): Gapdh and Actb, which were identified and validated as the most stable and suitable genes for RT-qPCR analysis in mouse testis development (48). 3β-hydroxysteroid dehydrogenase (Hsd3b1), a selective Leydig cell marker, has been used as a normalization factor for the analysis of genes expressed in Leydig cells (49). Data were analyzed using the 2−ΔΔCt method (50).

Western blot

Protein extraction

Testes were homogenized in ice-cold RIPA buffer (50 mM Tris HCl, 0.5% cholic acid, 0.1% SDS, 150 mM NaCl) containing a protease inhibitor cocktail (Sigma-Aldrich). For in vitro cultured tissues, the central necrotic area was carefully removed before protein extraction.

Determination of protein concentration by the Bradford method

Total protein concentration was measured in the homogenates. The assay was performed in a 96-well plate in a final volume of 200 µL containing the sample and Bradford’s solution (BioRad, Marnes-la-Coquette, France). After a 5-min incubation at RT, the optical density at 595 nm was measured using a Spark spectrophotometer (Tecan, Lyon, France).

Western blot

Protein extracts (20 μg) were separated on 12% resolving polyacrylamide gels (BioRad) and transferred onto nitrocellulose membranes (BioRad). The membranes were blocked with 5% milk or 5% BSA in TBST (0.2% Tween-20 in Tris Buffered Saline) for 30 min at 37°C before incubation overnight at 4°C with the appropriate primary antibody (S1 Table). After washes, membranes were incubated with the appropriate HRP-conjugated secondary antibodies (S1 Table). ECL reagents (BioRad) were used to detect the immunoreactivity. Images were captured with the ChemiDoc XRS+ Imaging System (BioRad). Densitometric analyses were performed using Image Lab™ v5.0 software (BioRad) and protein levels were normalized to β-actin or to the Leydig cell-specific marker 3β-HSD.

Intratesticular cholesterol levels

Total testicular lipids were extracted from ∼10 mg tissue with a Lipid Extraction kit (ab211044; Abcam, Paris, France) according to the manufacturer’s instructions. Dried lipid extracts were reconstituted in 200 µL of assay buffer. Intratesticular levels of total, free, and esterified cholesterol were measured with a colorimetric Cholesterol/Cholesteryl Ester assay kit (ab65359; Abcam) according to the manufacturer’s instructions.

Hormonal assays

Sample preparation for liquid chromatography coupled to mass spectrometry (LC-MS/MS)

Testicular fragments were homogenized in 200 µL (6 dpp tissues and in vitro matured tissues) or 400 µL (22 dpp and 36 dpp tissues) of 0.1 M phosphate buffer pH 7.4 with a cocktail of protease inhibitors (Sigma-Aldrich). Total protein concentration was measured as described above.

LC-MS/MS

The standards for androstenedione, testosterone, dehydroepiandrosterone (DHEA) and their stable labeled isotopes were obtained from Merck. Working solutions were prepared in methanol. Serial dilutions from working solutions were used to prepare seven-point calibration curves and 3 quality control levels for all analytes. For the final calibration and quality control solutions, PBS was used for testicular fragments and α-MEM for medium. The linearity ranges were from 0.05 to 10 ng/mL for androstenedione and testosterone, and from 1 to 200 ng/mL for DHEA. A simple deproteinization was carried out by automated sample preparation system, the CLAM-2030 (Shimadzu Corporation, Marne-la-Vallée, France) coupled to a 2D-UHPLC-MS/MS system. Once the sample on board, 30 µL was automatically pipetted in a pre-conditioned tube containing a filter, in which reagents were added, then mixed and filtered. Briefly, the polytetrafluoroethylene (PTFE) filter vial (0.45 μm pore size) was previously conditioned with 20 μL of methanol (Carlo Erba, Val-de-Reuil, France). Successively, 30 μL of sample and 60 μL of a mixture of isotopically labeled internal standards in acetonitrile were added. The mixture was agitated for 120 s (1900 rpm), then filtered by application of vacuum pressure (-60 to -65 kPa) for 120 s into a collection vial. Finally, 30 µL of the extract was injected in the 2D-UHPLC-MS/MS system.

Analysis was performed on a two dimensions ultra-performance liquid chromatograph-tandem mass spectrometer (2D-UHPLC-MS/MS) consisting of the following Shimadzu® modules (Shimadzu Corporation): an isocratic pump LC20AD SP, for pre-treatment mode, a binary pump consisting of coupling two isocratic pumps Nexera LC30AD for the analytical mode, an automated sampler SIL-30AC, a column oven CTO-20AC and a triple-quadrupole mass spectrometer LCMS-8060. The assay was broken down into two stages. The first was the pre-treatment where the sample was loaded on the perfusion column. The second step was the elution of the compounds of interest to the analytical column. The deproteinized extract performed by the CLAM-2030 was automatically transferred to the automated sampler, where 30 µL was directly analyzed into the chromatographic system. The LC-integrated online sample clean-up was performed using a perfusion column Shimadzu® MAYI-ODS (5 mm L × 2 mm I.D.). The first step consisted of loading the extract on the perfusion column with a mobile phase composed of 10 mM ammonium formate in water (Carlo Erba) at a flow rate of 0.5 mL/min during 2 min. Then the system switched to the analytical step to elute the analytes from the perfusion column to the analytical column to achieve chromatographic separation. During this step, the loading line was washed with propan-2-ol (Carlo Erba) during 3 min. Chromatographic separation was achieved on a Restek® Raptor Biphenyl (50 mm L × 3 mm I.D., 2.7 μm) maintained at 40 °C and a gradient of (A) 1 mM ammonium fluoride buffer in water (Carlo Erba) and (B) methanol (Carlo Erba) at a flow rate of 0.675 mL/min as follows: 0.0–2.10 min, 5% (B); 2.1–3.0 min, 5 to 65% (B); 3.0–4.75 min, 65% (B); 4.75–5.0 min, 65 to 70% (B); 5.0–6.6 min, 70% (B); 6.6–8.0, 70 to 75% (B); 8.0–8.5 min, 75 to 100% (B); 8.50–9.5 min, 100% (B); 9.5–9.6 min, 100% to 5% (B); 9.6–12.0 min, 5% (B). Detection and quantification were performed by scheduled-Multiple Reaction Monitoring (MRM) using 1 ms pause time and 50 ms dwell times to achieve sufficient points per peak. The interface parameters and common settings were as follows: interface voltage: 1 kV; nebulizing gas flow: 3 L/min; heating gas flow: 10 L/min; drying gas flow: 10 L/min; interface temperature: 400 °C; desolvation line (DL) temperature: 150 °C; heat block temperature: 500 °C; collision gas pressure 300 kPa. Compound-specific MRM parameters are shown in Supplementary File 1c.

Enzyme-linked immunosorbent assay (ELISA)

Steroids were extracted from testicular homogenates and from organotypic culture media with 5 volumes of diethyl ether (Carlo Erba). The organic phase was then recovered and evaporated in a water bath at 37°C. The tubes were stored at -20°C until use. Progesterone and estradiol levels were measured in 50 µL of testicular homogenates and in 50 µL of culture media using Cayman ELISA kit (582601 for progesterone and 501890 for estradiol, Cayman Chemical Company, Ann Arbor, MI, USA), according to the supplier’s recommendations. The sensitivity limit for the progesterone assay was 10 pg/mL with average intra- and inter-trial variations of 13.9% and 9.6%, respectively. The sensitivity limit for the estradiol assay was 20 pg/mL and the lower limit of detection was 6 pg/mL with average intra- and inter-trial variations of 20.25% and 14.975%, respectively.

Aromatase activity

After homogenization of ∼10-50 mg of tissues in 500 μL of Aromatase Assay Buffer, the aromatase activity was measured with the fluorometric Aromatase (CYP19A) Activity assay kit (ab273306; Abcam) according to the manufacturer’s instructions.

Radioimmunoassay (RIA)

INSL3 was assayed by using RIA kit with 125I as tracer (RK-035-27, Phoenix Pharmaceuticals, Strasbourg, France). According to the manufacturer’s instructions, analyses were performed on 100 μL of culture media or tissue homogenates. Assay validation was assessed by determining the recovery of expected amounts of INSL3 in samples to which exogenous INSL3 was added. The sensitivity of the INSL3 RIA kit was 60.4 pg/mL and the range of the assay was 10-1280 pg/mL.

Statistical analyses

Statistical analyses were carried out with GraphPad Prism 8 software (GraphPad Software Inc., La Jolla, CA, USA). Data are presented as means ± SEM. The non-parametric Mann-Whitney test was used to compare in vitro cultures and in vivo controls (D16 FT vs 22 dpp, D16 CSF vs 22 dpp, D30 FT vs 36 dpp, D30 CSF vs 36 dpp, D30 FT + hCG vs 36 dpp, D30 CSF + hCG vs 36 dpp); cultures of fresh and CSF tissues (6 dpp vs 6 dpp CSF, D16 FT vs D16 CSF, D30 FT vs D30 CSF, D30 FT + hCG vs D30 CSF + hCG); cultures with or without hCG (D30 FT + hCG vs D30 FT, D30 CSF + hCG vs D30 CSF). A value of P < 0.05 was considered statistically significant.

Supporting information